# Part 2: Predicting Housing Prices in Cook County



## Introduction

In Part 1, we performed some basic Exploratory Data Analysis (EDA), laying out the thought process that may lead to certain modeling decisions. Then, we added a few new features to the dataset and cleaned the data in the process.

In this part, we will specify and fit a linear model to a few features of the housing data to predict house prices. Then, we will analyze the error of the model and brainstorm ways to improve the model's performance. 

## The CCAO Dataset

We'll work with the dataset from the Cook County Assessor's Office (CCAO) in Illinois. This government institution determines property taxes across most of Chicago's metropolitan areas and nearby suburbs. In the United States, all property owners must pay property taxes, which are then used to fund public services, including education, road maintenance, and sanitation. These property tax assessments are based on property values estimated using statistical models considering multiple factors, such as real estate value and construction cost.

While our simple model explains some of the variability in price, there is certainly still a lot of room for improvement —— one reason is we have been only using 1 or 2 features (out of a total of 70+)! In the next part, you will engineer and incorporate more features to improve the model's fairness and accuracy.


# Building the Model

It is time to build our model. Next, we will conduct feature engineering on your training data using the `feature_engine_final` function, and fit and evaluate the model with the training data.


```python
# Import all the necessary libraries

import numpy as np
import pandas as pd
from pandas.api.types import CategoricalDtype

%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
from sklearn import linear_model as lm
from sklearn.model_selection import KFold

import warnings
warnings.filterwarnings("ignore")

import zipfile
import os

from sklearn.preprocessing import OneHotEncoder
```

```python
training_data = pd.read_csv('cook_county_train.csv')

```



## Helper Functions

 

```python
# Define additional helper functions
def rmse(predicted, actual):
    return np.sqrt(np.mean((actual - predicted)**2))


def get_outlier_bounds(data, column):
    quartiles = data[column].quantile([0.25, 0.75])
    IQR = quartiles[0.75] - quartiles[0.25]
    return [quartiles[0.25] - 1.5 * IQR, quartiles[0.75] + 5 * IQR]

def remove_outliers(data, variable, lower=-np.inf, upper=np.inf):
    return data[(data[variable] > lower) & (data[variable] < upper)]

def add_total_bathrooms(data):
    with_rooms = data.copy()
    with_rooms['Bathrooms'] = with_rooms['Description'].str.extract(r'(\d+\.\d+) of which are bathrooms').astype('float')
    return with_rooms


def good_condition(data):
    data['in_good_condition'] = data['Repair Condition'].apply(lambda x: 1 if x == 3 else 0)
    return data

def good_basement(data):
    data['has_good_basement'] = data['Basement'].apply(lambda x: 1 if x in [1, 3] else 0)
    return data

def add_indicators(data, column):
    ohe = OneHotEncoder(drop = 'first')
    ohe.fit(data[[column]])
    ohe_columns = pd.DataFrame(ohe.transform(data[[column]]).toarray(), columns = ohe.get_feature_names_out(), index = data.index)
    return data.merge(ohe_columns, left_index=True, right_index=True)


```


##  Pipeline Function


The function `feature_engine_final` will be our pipeline for feature engineering. It will return X, the feature matrix of predictors, and Y, the outcome for the training data. It will also allow the parameter `is_test_set` to be set to true to return only X, the feature matrix for an unlabeled test dataset so the model can be evaluated on unseen data elsewhere.

```python
def feature_engine_final(data, is_test_set=False):
    if not is_test_set:
        # Processing for the training set (i.e. not the test set)
        # Involves references to sale price
        # Involves filtering certain rows or removing outliers
        data['Log Sale Price'] = np.log(data['Sale Price'])
        data = remove_outliers(data, 'Sale Price', lower = 499)
        data = remove_outliers(data, 'Sale Price', upper = 10000000)


    # Processing for both test and training set
    # Does not involve references to sale price
    # Does not involve removing any rows

    data = add_total_bathrooms(data) #Add bathrooms column
    data['Log Building Square Feet'] = np.log(data['Building Square Feet']) #Log transform building square ft
    data = add_indicators(data, 'Basement')
    data = add_indicators(data, 'Repair Condition')
    data = add_indicators(data, 'Basement Finish')
    data = add_indicators(data, 'Roof Material')
    data = add_indicators(data, 'Property Class')

    data['Estimate'] = data['Estimate (Land)'] + data['Estimate (Building)']
    data['Adj Estimate'] = data['Estimate'] + 0.001
    data['Log Estimate'] = np.log(data['Adj Estimate'])
    

    # Return predictors (X) and response (Y) variables separately
    if is_test_set:
        # Predictors 
        X = data[['Bathrooms',
                  'Log Building Square Feet',
                  'Fireplaces',
                  'Central Air',
                  'Basement_2.0',
                  'Basement_3.0', 
                  'Basement_4.0',
                  'Repair Condition_2.0',
                  'Repair Condition_3.0',
                  'Basement Finish_3.0',
                  'Roof Material_2.0',
                  'Roof Material_3.0',
                  'Roof Material_4.0',
                  'Roof Material_5.0',
                  'Roof Material_6.0',
                  'Log Estimate'
                 ]]
        return X
    else:
        # Predictors. X will not include Log Sale Price
        X = data[['Bathrooms',
                  'Log Building Square Feet',
                  'Fireplaces',
                  'Central Air',
                  'Basement_2.0',
                  'Basement_3.0', 
                  'Basement_4.0',
                  'Repair Condition_2.0',
                  'Repair Condition_3.0',
                  'Basement Finish_3.0',
                  'Roof Material_2.0',
                  'Roof Material_3.0',
                  'Roof Material_4.0',
                  'Roof Material_5.0',
                  'Roof Material_6.0',
                  'Log Estimate'
                 ]]
        # Response variable
        Y = data['Log Sale Price']
        
        return X, Y
```


## Model Evaluation

To evaluate the model, we'll use 10-fold cross validation on our training data and compute RMSE. Since our model is set up to predict `Log Sale Price`, we'll get our predictions in log scale, then scale them back up to compute RMSE in terms of `Sale Price` as opposed to `Log Sale Price`.

```python
X, Y = feature_engine_final(training_data)

model = lm.LinearRegression(fit_intercept=True)

kf = KFold(n_splits=10)

cv_error = []
for train_idx, valid_idx in kf.split(X):
    # Split the data
    split_X_train, split_X_valid = X.iloc[train_idx], X.iloc[valid_idx]
    split_Y_train, split_Y_valid = Y.iloc[train_idx], Y.iloc[valid_idx]

    # Fit the model on the training split
    model.fit(split_X_train, split_Y_train)
    
    # Compute the RMSE on the validation split
    eval_df = pd.DataFrame({
        'Predicted': model.predict(split_X_valid),
        'Actual': split_Y_valid
    })
    
    eval_df['delog_predicted'] = np.exp(eval_df['Predicted'])
    eval_df['delog_actual'] = np.exp(eval_df['Actual'])

    cv_error.append(rmse(eval_df['delog_predicted'], eval_df['delog_actual']))

print(f'10-fold cross validated RMSE: ${np.round(np.mean(cv_error), 2)}')
```

Output:
```
10-fold cross validated RMSE: $183615.92
```

Tests for q1c (points: 1)
```
>>> q1c.lower() in ['a', 'b', 'c']
True
```


Tests for q2 (points: 2)
```
>>> bool(train.shape == (163833, 62))
True
```
```
>>> bool(validation.shape == (40959, 62))
True
```
```
>>> bool(np.isclose(train['Sale Price'].mean(), 244939.22668204817, atol=0.1))
True
```
```
>>> bool(np.allclose(validation.index[-5:], [153946, 117415, 9448, 188605, 3223]))
True
```
```
>>> bool(np.isclose(validation['Sale Price'].mean(), 246066.1821089382, atol=0.1))
True
```


Tests for q3a (points: 2)
```
>>> (q3a in ['>=', '=', '<=']) == True
True
```
```
>>> bool(q3a == '<=')
True
```


Tests for q3b (points: 3)
```
>>> isinstance(X_train_m1, pd.core.frame.DataFrame) and isinstance(Y_train_m1, pd.core.series.Series) and isinstance(X_valid_m1, pd.core.frame.DataFrame) and isinstance(Y_valid_m1, pd.core.series.Series) and isinstance(X_train_m2, pd.core.frame.DataFrame) and isinstance(Y_train_m2, pd.core.series.Series) and isinstance(X_valid_m2, pd.core.frame.DataFrame) and isinstance(Y_valid_m2, pd.core.series.Series)
True
```
```
>>> assert len(m2_pipelines) == 5
>>> assert log_transform in set([p[0] for p in m2_pipelines])
```


Tests for q3c (points: 2)
```
>>> bool(np.isclose(Y_fitted_m1.max(), 18.37391180106831, atol=0.0001) == True)
True
```
```
>>> bool(np.isclose(Y_fitted_m2.max(), 16.38031927932594, atol=0.0001) == True)
True
```
```
>>> bool(np.isclose(Y_predicted_m1.max(), 37.70679143176044, atol=0.0001) == True)
True
```
```
>>> bool(np.isclose(Y_predicted_m2.max(), 20.856723206185926, atol=0.0001) == True)
True
```


Tests for q4b (points: 1)
```
>>> bool(q4b.lower() in ['regressive', 'fair', 'progressive'])
True
```


Tests for q4ci (points: 1)
```
>>> len(compute_CV_error(X_train_m1, Y_train_m1, folds=4)) == 4
True
```
```
>>> len(compute_CV_error(X_train_m1, Y_train_m1, folds=6)) == 6
True
```


Tests for q4cii (points: 1)
```
>>> np.random.seed(1337)
>>> bool(np.isclose(sum(cv_m1), 3.3746790542844085, atol=0.0001))
True
```
```
>>> np.random.seed(1337)
>>> bool(np.isclose(sum(cv_m2), 3.190860696081888, atol=0.0001))
True
```


Tests for q5d (points: 3)
```
>>> bool(check_rmse_threshold(200000))
True
```
```
>>> bool(check_rmse_threshold(240000))
True
```
```
>>> bool(check_rmse_threshold(280000))
True
```
```
>>> check_rmse_threshold.signature == (feature_engine_final, 'cook_county_train.csv', None)
True
```


Tests for q6a (points: 1)
```
>>> bool(rmse_cheap >= 0)
True
```
```
>>> bool(rmse_expensive >= 0)
True
```
```
>>> bool(0 <= prop_overest_cheap <= 1)
True
```
```
>>> bool(0 <= prop_overest_expensive <= 1)
True
```


Tests for q6b (points: 2)
```
>>> bool(prop_overest_interval(preds_df, 10, 14) >= 0 and prop_overest_interval(preds_df, 10, 14) <= 1)
True
```
```
>>> bool(rmse_interval(preds_df, 10, 14) < 1000000000.0 and rmse_interval(preds_df, 10, 14) > 0)
True
```


Tests for qsurvey (points: 1)
```
>>> assert isinstance(qsurvey, str)
>>> assert ' ' not in qsurvey
>>> assert len(qsurvey) == 11
>>> assert qsurvey[0] == 'G' and qsurvey[-1] == 's'
```
